# DO NOT CREATE OR DELETE ANY CELLS
## DO NOT CREATE OR DELETE ANY CELLS
### DO NOT CREATE OR DELETE ANY CELLS
#### DO NOT CREATE OR DELETE ANY CELLS

# Objects and Data Structures Assessment Test

#### I Strongly advise you take the questions and practice on a different notebook and then when you are certain you have to correct code bring it in here and paste in the approriate cell.

### To answer the questions first delete this line `raise NotImplementedError()` and then write your code after the `#your code here` comment

## 1. Strings

- Create a variable **s** and put the string  **hello world** in it, then extract every second letter and store it in a variable called **letter_sequence**.
- Create a formatted string using **format** and **{}** and store it in variable called **greeting** such that If "Jadesola" is inserted as the name and "25" as the age, the result should be **Hello Jadesola, you are 25 years old**.

<!-- grade_id: cell-088358cd80c89256, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# creating a variable
s = "hello world"
print(s)

# Extracting every second letter
letter_sequence = s[::2]
print(letter_sequence)

# string formatting with .format() method
greeting = "Hello {}, you are {} years old"
greeting.format('Jadesola', 25)
```

Output:
```
hello world
hlowrd
'Hello Jadesola, you are 25 years old'
```

<!-- grade_id: cell-8c64ce6723dba25f, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
#do not modify this cell
#do not modify this cell
#do not modify this cell
assert letter_sequence == "hlowrd"
assert greeting.format("Jackson", 25) == "Hello Jackson, you are 25 years old"
```

## 2. Lists

- With a list defined like this `list3 = [1, 2, [3, 4, 'hello']]` modify it by:
    - Changing "hello" to "goodbye".
    - Append **100** to the list.
    - Removing **2** from the list.
    - Reverse the entire list.
    - Store the final result in a variable called **modiFied**

<!-- grade_id: cell-fac327c594288b28, points: 0, grade: false, solution: true -->
```python
list3 = [1, 2, [3, 4, 'hello']]
# YOUR CODE HERE
# changing "hello" to "goodbye"
list3[2][2] = "goodbye"
print(list3)

# appending 100 to the list
list3.append(100)
print(list3)

# removing 2 from the list
list3.remove(2)
print(list3)

# reverse the entire list and store in variable modiFied
modiFied = list3[::-1]
print(modiFied)
```

Output:
```
[1, 2, [3, 4, 'goodbye']]
[1, 2, [3, 4, 'goodbye'], 100]
[1, [3, 4, 'goodbye'], 100]
[100, [3, 4, 'goodbye'], 1]
```

<!-- grade_id: cell-97260ca982a851e2, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
#do not modify this cell
#do not modify this cell
#do not modify this cell
assert modiFied == [
    10 ** 2,
    [9/3, 16**0.5, "good" + "bye"],
    2 - 1
]
```

## 3. Dictionaries

Given the dictionary:
`d = {"k1": [{"nest_key": ["this is deep", ["hello"]]}], "age": 30}`
- Extract "hello" and store it in **dict_answer**.
- Add a new key **name** with the value **Alice**.
- Increase **age** by **5**.

<!-- grade_id: cell-ea72189458dfbb55, points: 0, grade: false, solution: true -->
```python
d = {"k1": [{"nest_key": ["this is deep", ["hello"]]}], "age": 30}
# YOUR CODE HERE
# Extract "hello" from dict
dict_answer = d["k1"][0]["nest_key"][1][0]
print(dict_answer)

# Add a new key "name" with the value "Alice"
d["name"] = "Alice"
print(d)

# Increase age by 5
d["age"] += 5
print(d)

```

Output:
```
hello
{'k1': [{'nest_key': ['this is deep', ['hello']]}], 'age': 30, 'name': 'Alice'}
{'k1': [{'nest_key': ['this is deep', ['hello']]}], 'age': 35, 'name': 'Alice'}
```

<!-- grade_id: cell-92c0ac6f156e2a56, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
#do not modify this cell
#do not modify this cell
#do not modify this cell
assert dict_answer == "hello"
assert d["name"] == "Alice"
assert d["age"] == 35
```

## 4. Tuples

- Create a tuple **my_tuple** with values (5, 15, 3).
- Unpack it into three variables: a, b, c.
- Concatenate my_tuple with (4, 5, 6) and store it in new_tuple.

<!-- grade_id: cell-e778568dc2022c81, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# create a tuple
my_tuple = (5, 15, 3)
print(my_tuple)

# Unpack it into three variables: a, b, c.
a, b, c = my_tuple
print(a, b, c)

# Concatenate my_tuple with (4, 5, 6)
new_tuple = my_tuple + (4, 5, 6)
print(new_tuple)
```

Output:
```
(5, 15, 3)
5 15 3
(5, 15, 3, 4, 5, 6)
```

<!-- grade_id: cell-c84a7e6cfef7d56b, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
#do not modify this cell
#do not modify this cell
#do not modify this cell
assert a == 15/3 and b == 7+8 and c == 9**0.5
```

## 5. Sets

Use a set to find the unique values of the list below and assign it to a variable **unique_values**:

<!-- grade_id: cell-9f915e6b6eb8cd39, points: 0, grade: false, solution: true -->
```python
list5 = [1,2,2,33,4,4,11,22,3,3,2]
# YOUR CODE HERE
# finding unique values from list5
unique_values = set(list5)
print(unique_values)
```

Output:
```
{1, 2, 33, 4, 3, 11, 22}
```

<!-- grade_id: cell-9c9f34983f081b5a, points: 2, grade: true, solution: false -->
```python
#do not modify this cell
#do not modify this cell
#do not modify this cell
#do not modify this cell
assert unique_values ==  {1, 2, 3, 4, 11, 22, 33}
```

## Great Job on your first assessment!